# Working with a database

This week's assignment has a few basic steps.  First, we're going to pull some data down off the internet and store it into our MySQL database.  Make sure that you use your username as part of the table name as show in the examples so that you don't create a problem for other students.

Then, we'll merge that with some data already in the database and calculate a few results.  When it comes to calculating the results, you can do so either with SQL or with Pandas operations.


## PART 1: Setup your database connection and table name

In the code below, change the value of the variable `MYTABLE` to use your own username rather that `'pboal'`

You can then use `MYTABLE` in the rest of your code to reference that table name.

```python
import os
import pymysql
from sqlalchemy import create_engine
import pandas as pd
import numpy as np
```

```python
host = 'slucor2022-instance-1.cgdcoitnku0k.us-east-1.rds.amazonaws.com'
port = '3306'
user = 'slucor2022'
password = 'SLUcor2022'
database = 'hds5210'
```

```python
conn = create_engine('mysql+pymysql://' + 
                     user + ':' + 
                     password + '@' + 
                     host + '/' + 
                     database, echo=False)
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
import getpass
myname = getpass.getuser().split('-')[1]
myname
MYTABLE=myname + '_data'


```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
assert(MYTABLE != 'pboal')
assert(conn.connect())
```

## PART 2: Bring in outside data

Grab the data from this URL and put it into a database table named with your `username_data`.

https://opendata.arcgis.com/api/v3/datasets/ed0bc689bde246b18835c7529ba4efb4_0/downloads/data?format=csv&spatialRefId=4326

By the end of your cell, the table should be created.  The tests are going to verify that the table exists and looks right.

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
ehr_df = pd.read_csv("../EHR.csv")
```

```python
ehr_df.to_sql(MYTABLE, conn, if_exists = 'replace')
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
assert(dxyz.shape == (323,22))
assert(list(dxyz.columns) == ['index', 'X', 'Y', 'OBJECTID', 'Provider_Name', 'NPI', 'CCN',
       'Business_Street_Address', 'Business_City', 'Business_County',
       'Business_ZIP_Code', 'Business_State_Territory', 'Payment_Year_Number',
       'Program_Type', 'Medicaid_EP_Hospital_Type', 'total_payments',
       'Last_Payment_Criteria', 'Recent_Disbursement_Amount', 'Latitude',
       'Longitude', 'Last_Program_Year', 'Last_Payment_Year'])
```

## PART 3: Combine with other data in the database

In the database, there is an existing table called `population`.  We want to merge the DHCS datafile loaded above with the population data available in this other database table  The tables can be merged on `MYTABLE`'s `Business_ZIP_Code` field and `population`'s `Zip` field.

Note that not all `Zip_Codes` from your downloaded file have to be in the `population` table.  If they aren't, then I want you to eliminate the non-matching records.  That is, only keep the records that have a matching ZIP code in both sets of data.

Answer the question:
Which providers are located in the zipcode with the largest population?

Put your answer in the form `answer = ['a', 'list', 'of', 'NPI', 'like', '1593042103]`

```python
ans = pd.read_sql_query("""
  SELECT h.NPI
  FROM 
    svelicheti_data h JOIN
    population s ON h.Business_ZIP_Code = s.Zip
  ORDER BY Population DESC
  LIMIT 1
""", conn)
```

```python
answer = list(ans['NPI'])
answer[0] = str(answer[0])
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == list)
assert(answer == ['1194016923'])

```

## PART 4: Total by hospital type

This is a multistep process:
* From our downloaded data file, compute the `Total payments` per ZIP code and Medicaid EP Hospital Type.
* Then merge that with the `population` data to compute a `Total payments` per person.
* Then average that across all of the `Medicaid EP Hospital Types` to get an average per persona payment for these type of hospital.

Your answer should be in structure of a data frame with at least two columns:
* Medicaid_EP_Hospital_Type
* Avg_Pay_per_Capita

## First using SQL

```python
computed_tot_pay = pd.read_sql_query("""
    SELECT Business_ZIP_Code, Medicaid_EP_Hospital_Type, Sum(total_payments) as tot_payments
  FROM 
    svelicheti_data
  GROUP BY Business_ZIP_Code, Medicaid_EP_Hospital_Type
""", conn)
computed_tot_pay.to_sql('comp_tot_pay', conn, if_exists = 'replace')
```

```python
computed_tot_pay = pd.read_sql_query("""
    SELECT Medicaid_EP_Hospital_Type, AVG(tot_payments/Population) as Avg_Pay_per_Capita
  FROM 
    population JOIN
    comp_tot_pay ON Zip = Business_ZIP_Code
  GROUP BY Medicaid_EP_Hospital_Type
""", conn)
answer = computed_tot_pay
```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 106.583)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 329.551)

```

## Now using Pandas instead of SQL

```python
population_df = pd.read_sql_query("""SELECT * FROM population""", conn)
```

```python
computed_tot_pay = pd.pivot_table(ehr_df, index = ['Business_ZIP_Code', 'Medicaid_EP_Hospital_Type'], values = 'total_payments', aggfunc = np.sum)
computed_tot_pay = computed_tot_pay.reset_index()
```

```python
merged_data = computed_tot_pay.merge(population_df, how='inner', left_on = 'Business_ZIP_Code', right_on ='Zip')
```

```python
merged_data['tot_pay_per_person'] = merged_data['total_payments'] / merged_data['Population']  
```

```python
answer = merged_data.groupby(['Medicaid_EP_Hospital_Type'])['tot_pay_per_person'].agg(Avg_Pay_per_Capita = ('mean'))
answer = answer.reset_index()
```

```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 106.583)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 329.551)
```



---



## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Follow the instruction on the prompt below to either ssave and submit your work, or continue working.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

---

```python
a=input('''
Are you ready to submit your work?
1. Click the Save icon (or do Ctrl-S / Cmd-S)
2. Type "yes" or "no" below
3. Press Enter

''')

if a=='yes':
    !git pull
    !git add week15_assignment.ipynb
    !git commit -a -m "Submitting the week 15 programming assignment"
    !git push
else:
    print('''
    
OK. We can wait.
''')
```

Output:
```
Are you ready to submit your work?
1. Click the Save icon (or do Ctrl-S / Cmd-S)
2. Type "yes" or "no" below
3. Press Enter

 yes
Already up to date.
On branch main
Your branch is up to date with 'origin/main'.

Untracked files:
	[31m../EHR.csv[m
	[31m../HospInfo.csv[m
	[31m../Untitled.ipynb[m
	[31m../Untitled1.ipynb[m
	[31m../apache_patients.csv[m
	[31m../week 03/week03_class.ipynb[m
	[31m../week 04/week04_examples.ipynb[m
	[31m../week 04/week04_lookups.ipynb[m
	[31m../week 05/[m
	[31m../week 06/module29-xml-examples.ipynb[m
	[31m../week 06/week06_inclass_1.ipynb[m
	[31m../week 10/Number_of_Cancer_Cases_for_All_Cancer_Sites_by_Jurisdiction__Gender__and_Race__Maryland_2009.csv[m
	[31m../week 13/HospInfo.csv[m
	[31m../week02/[m

nothing added to commit but untracked files present
Everything up-to-date
```

```python

```